```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

<table style="width: 100%;" id="nb-header">
    <tr style="background-color: transparent;"><td>
        <img src="https://data-88e.github.io/assets/images/blue_text.png" width="250px" style="margin-left: 0;" />
    </td><td>
        <p style="text-align: right; font-size: 10pt;"><strong>Economic Models</strong>, Data 88E<br>
            Dr. Eric Van Dusen <br>
        Alan Liang <br>
        Amal Bhatnagar <br>
        Sreeja Apparaju <br>
        Akhil Venkatesh</p></td></tr>
</table>

# Lab 6: Modeling Income Inequality & Development

```python
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
from sympy import * 
from scipy.optimize import minimize
import ipywidgets as widgets
from datascience import *
from ipywidgets import interact, interactive, fixed, interact_manual, widgets
from IPython.display import display
import pandas as pd
%matplotlib inline 

def find_k(lnx, lny):
    """Determines the best value for k in the simplified Rao and Tam model.
    
    :param lnx: log of the income percentiles (out of 1)
    :param lny: log of the income shares (out of 1)
    :returns: the best value for k.
    """
    def obj(k):
        return np.mean(np.abs(lny - k*lnx))
    res = minimize(obj, 2)
    return res
```

## Part 1: Inequality

## The Lorenz Curve
There are many mathematical models used to model the Lorenz curve. For this homework, we will work with the simplified Rao-Tam Lorenz curve<sup>1</sup>, as defined by:
$$L(x)=x^k,\quad \text{where} \; k>1$$

<!-- END QUESTION -->

### Fitting the Lorenz Curve to Data
Now, let's try to fit this model to the existing US data. From the 2018 census data on *Income and Poverty in the United States*<sup>2</sup> published by the US census, we get:

```python
us = Table().read_table("us_income.csv")
us
```

For each row of the table, `income_share` shows the proportion of total income an `income_group` owns. For example, the richest 20% (5th quintile) earned 51.9% of the total national income.

**Question 1.1:**
Let's convert this table to points on the Lorenz curve. Recall that the Lorenz curve's y-axis is the cumulative income share: at the 20th percentile, the total income share is 3.1%, while at the 40th percentile, the total income share is 8.3% + 3.1% = 11.4%. 

*Hint: use `np.cumsum`.*


```python
income_group = ...
cum_income_share = ...

us_cumulative = Table().with_columns("income_percentile", income_group, "cum_income_share", cum_income_share)
us_cumulative
```

```python
grader.check("q1_1")
```

**Question 1.2:**
In addition, the US census bureau reported that the top 5 percent of income earners' income share was 23.1%. What is the corresponding $(x,y)$ point to this data point? 


```python
x_val = ...
y_val = ...
```

```python
grader.check("q1_2")
```

**Question 1.3:**
Add this pair of coordinates to our table, by appending `income_group` and `cum_income_share` with `x` and `y`.


```python
income_group_with_x = ...
cum_income_share_with_y = ...

us_cumulative_extended = Table().with_columns(
    "income_percentile", income_group_with_x, 
    "cum_income_share", cum_income_share_with_y)
us_cumulative_extended
```

```python
grader.check("q1_3")
```

**Question 1.4:**
To more easily fit our data, we will take the natural log of the curve to make it linear. Which one of the equations below equals the natural log of the curve?
<div>
a. $$\ln{y} = k + \ln{x}$$
<br>
b. $$\ln{x} = k + \ln{y}$$
<br>
c. $$\ln{y} = k \ln{x}$$
<br>
d. $$\ln{x} = k \ln{y}$$
<br>
</div>

Assign the letter corresponding to your answer to `q2_4` below.


```python
q1_4 = ...
```

```python
grader.check("q1_4")
```

**Question 1.5:**
Transform the data and append it to the new table. Use the [`np.log`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.log.html) function.


```python
log_x = ...
log_y = ...
us_log = Table().with_columns("log_income_percentile", log_x, "log_cum_income_share", log_y)
us_log
```

```python
grader.check("q1_5")
```

The following dataset come from a randomized controlled trial household survey carried out in Eastern Kenya in 2007-2008. 

The purpose of the study was to understand how to promote the use of WaterGuard, a dilute sodium hypochlorite solution that was promoted for Point-of-use household water disinfection.  There were seven arms in the study, which will be more fully described in the following chart:


## Part 2: Development


<img src="Slide1.png"  />

Within this table you can see the seven treatments arms -  control plus three treatments -  in the bolded boxes in the middle with the number of springs and households. The study was carried out as a part of a study of households who gather drinking water from springs in a rural area.  The three boxes at the bottom describe the three rounds of data collection - a baseline before the treatment, and a short term and long term follow-up.  

<!-- **Notebook Outline**

1. [Mapping](#Mapping)
2. [Balance Check](#Balance)
3. [Baseline and a Randomly Selected Compound](#Baseline)
4. [Chlorine Usage outcome variables](#Chlorine)
5. [Graph of outcomes by Treatment Arm](#Graph)  -->

```python
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
import pandas as pd
from pandas import read_stata
from ipyleaflet import Map, basemaps, Marker, AwesomeIcon
```

## Mapping

<div id="Mapping"></div>


This first section works with a package in Jupyter called ipyleaflet.

`ipyleaflet`;
the documentation is [here](https://ipyleaflet.readthedocs.io/en/latest/)
and it is worth a short read through if you are interested.


We want to use two different base maps - one is a satellite layer and oen is the Open Street Map layer.  

We will start by reading in a dataset of the coordinates of the springs that are used in the WaterGuard Promotion (WGP) study.  These springs were randomized into seven different treatment arms.  The springs are identified by a unique numerical id tag, and the common name in the local language.  


```python
springsGPS = Table.read_table('WGPgps_forData8.csv')
springsGPS
```

```python
# make a table wth just the North and East Gps columns 
locations = springsGPS.select("gpsn1", "gpse1")
locations
```

Where in the world are we?

First of all, let's look at the mean for the latitude and longitude, and we can center our map there:

```python
mean_longitude = springsGPS.column('gpse1').mean()
mean_latitude = springsGPS.column('gpsn1').mean()

print("Mean of 'gpse1':", mean_longitude)
print("Mean of 'gpsn1':", mean_latitude)
```

The code cell below should display a map. However, it may not run the first time you click it - if this happens, try running all the cells above this one and then refreshing your browser. After a few refreshes, the maps should load.

`ipyleaflet` can take the following parameters: 
 - `center` - where to place the center of the map (you can still zoom in and out) 
 - `zoom` - zoom level to start at 
 - `basemap` - what tiles to draw the underlying, base map from 
 - `layout` - specifying the size that it will display 

```python
center = [0.4, 34.4]
zoom = 12
basemap=basemaps.Esri.WorldImagery
basemap_osm=basemaps.OpenStreetMap.Mapnik
layout={'width': '800px', 'height': '600px'}

Map(basemap=basemap, center=center, zoom=zoom, layout=layout)
```

Let's make a map of our sample sites (springs):

```python
m = Map(basemap=basemap, center=center, zoom=zoom, layout=layout)

# Add the rows from the dataset
for row in springsGPS.rows:
    latitude = row.item('gpsn1')
    longitude = row.item('gpse1')
    marker = Marker(location=(latitude, longitude))
    m.add_layer(marker)

m
```

The most interesting bit of data, the Treatment Arm, is still not being used. Let's assign different colors to the different treatment arms so that when we map it we can see if the arms appear to be randomly distributed.

The following is function assigns the 7 different treatment arms to a set of colors. [Here](https://www.w3.org/TR/css3-color/#html4) is the colors reference if you are interested!

```python
def color(arm):
    if arm == 1:
        return 'black'
    elif arm == 2:
        return 'red'
    elif arm == 3:
        return 'purple'
    elif arm == 4:
        return 'green'
    elif arm == 5:
        return 'blue'
    elif arm == 6:
        return 'pink'
    elif arm == 7:
        return 'orange'
```

```python
# Using the .apply method, you can apply any function to a data frame
colors = springsGPS.apply(color, "treatment_arm")
springsGPS = springsGPS.with_column("color", colors)
springsGPS
```

## Map with markers that have a different color for each treatment arm

```python
m = Map(basemap=basemap, center=center, zoom=zoom, layout=layout)

for row in springsGPS.rows:
    latitude = row.item('gpsn1')
    longitude = row.item('gpse1')
    color = row.item('color')
    
    marker = Marker(
        location=(latitude, longitude),
        draggable=False,  # Set to True if you want to make the markers draggable
        title=color,      # Set the marker title to the color for tooltip
        alt=color         # Set the alt text to the color
    )
    marker.icon = AwesomeIcon(name='circle', marker_color=color)
    
    m.add_layer(marker)

m
```

```python
#change the Basemap layer to Open Street Map

m=Map(basemap=basemap_osm, center=center, zoom=zoom, layout=layout)

# Iterate through the rows in the dataset
for row in springsGPS.rows:
    latitude = row.item('gpsn1')
    longitude = row.item('gpse1')
    color = row.item('color')
    
    # Create a marker with the specified color
    marker = Marker(
        location=(latitude, longitude),
        draggable=False,  # Set to True if you want to make the markers draggable
        title=color,      # Set the marker title to the color for tooltip
        alt=color         # Set the alt text to the color
    )
    
    # Apply the specified color to the marker
    marker.icon = AwesomeIcon(name='circle', marker_color=color)
    
    m.add_layer(marker)

m
```

Do the colors seem randomly distributed?

In fact, the randomization was performed on just a list of the springs using a random number generator. 
It did not take spatial distribution into effect.  


**Non-Graded Reflection Question:**
 A Thought Experiment on Spatial Randomization
- What could you do to test whether the Treatment arms are spatially distributed?
- What could you do to randomize the treatment arms over space?

<div id="Balance"></div>

## Balance Check and Variable Names

### Baseline Survey 
This is our first look at the survey dataset.  These are a limited set of questions/answers from a simple and short baseline survey. However it is a lot bigger and messier than the datasets we have seen so far and in Data 8. 

Data variable names follow along with the survey below, referred to by the section, a,b,c... number, 1,2,3... and a few words about the question. 

The purpose of this section will be: 
* to get a familiarity with the dataset, 
* to look at some background descriptor variables of the households, 
* to start to think about missing values and coding of subsets of the data.  
* to check the randomization of households by seeeing if the different arms of the study are balanced across some of the key baseline variables.  

**The surveys that illustrate the raw data names are in a file linked [here](https://drive.google.com/open?id=1UVoiVn7LJ4rn7WEb-9BJ96jmdJ2FBk60). You have to go and look through this survey to understand the variables.**

**The code sheet that has the codes for some of the possible answers are in a file linked [here](https://drive.google.com/file/d/1iinJXExeVKV4Dm7tRKOiotoYUDSXMyqc). You have to go and look through this code sheet in a later section.**

```python
WGP_baseline = Table.read_table("WGP_baseline_Data8.csv")
WGP_baseline
```

### Misssing values 

If you look through the dataset above, and scroll to the right a ways to some of the last variables, you will notice that that there are a lot of cells with NaN, which means a missing value. For these cells no data was entered at the time of data entry. In some cases it may be appropriate to enter a zero and carry on with the analyis.  



```python
WGP_base_dfna = WGP_baseline.to_df().fillna(0)
WGP_table = Table.from_df(WGP_base_dfna)
WGP_table
```

 Look at the variable names, and then look at the survey form to find the concordance of codes

```python
# Here is a list of all of the possible categories / columns
list(WGP_table)
```

### What are some Variables that we want to specifically look at? ###

There are a lot of variables here and it can be kind of overwhelming, but it is good to see how many columns there can be in a comprehensive survey dataset.  

#### Front Page information - A variables

- household id
- spring id
- interviewer id

#### Information about respondent - B variables 

- tribe
- education
- age
- gender 
- group membership

#### Water Guard Use - C variables

For Waterguard (WG) usage

- `c1a` - Whether the respondent has ever heard of WG
- `c2a` - Whether the respondent has ever used WG
- `c3a` - Whether the respondent's water is currently treated with WG
- `c4a` - Whether the respondent has used WG in the past month

#### Durable / Capital Goods - D variables

- Whether the respondent has electricity / latrine / iron roof
- Number of of bicycle / radio / hoe / beds owned
- Number of animals owned

#### Child Health - E variables

- `e1_num_kids_under_5`: Number of kids under 5
- `e2_`:  This table becomes tricky because it has a different format. Each kid in the table is numbered 01, 02 and so on, and then the subsequent questions are keyed to that child number. e.g. `e2e_01_d_diarrhea`, `e2e_02_d_diarrhea` represent whether child 1 and 2 respectively have diarrhea. In total, four diseases are recorded:
    - Cough
    - Diarrhea
    - Malaria
    - Vomiting



 

### The Treatment Arm 

In the study, arm 1 is control, while Arms 2-7 are different types of treatment interventions:
 
- Arm 1 - Control
- Arm 2 - Household Script
- Arm 3 - Community Script
- Arm 4 - HH + Community Script
- Arm 5 - Flat-Fee Promoter + Coupons
- Arm 6 - Incentivized Promoter + Coupons
- Arm 7 - Incentivized Promoter + Dispenser at Spring

Let's check how many households are in each treatment arm.


```python
WGP_table.group("treatment_arm")
```

### Baseline Check - Exposure to Water Guard Use 

Let's see how many households have ever used Water Guards.

The data is currently Coded as 1 = Yes and 2 = No, so we can't really make sense of the Mean of the variable in its current form. Instead, we will make a new column/variable with the 1 or 2 answers translated into Yes or No.
Notably, we must first filter out respondents that had missing values (with value 0) for this question.



```python
WGP_ever = WGP_table.where('c2a_wg_used_ever', are.above(0))
WGP_ever.group("c2a_wg_used_ever")
```

```python
#This helper function goes through a column of choice, and spits out yes or no based off each value in the column. It returns an array of these yes and no's
def translate_to_yesno(table, col):
    dummy=[]
    table=table.where(col, are.above(0))
    for i in np.arange(table.num_rows):
        if table.column(col).item(i) == 1:
            dummy.append('Yes')
        else: #if not 1 then its 2 and 2 means no
            dummy.append("No")
    return dummy
```

```python
new = translate_to_yesno(WGP_ever, 'c2a_wg_used_ever')
WGP_ever = WGP_ever.with_column('c2a_wg_used_ever',new)
WGP_ever.group('c2a_wg_used_ever')
```

### Pivoting and Balance Checks

Now we will use a command called **Pivot** to create a new table that has the percent of households who have ever used Water Guard within each Treatment Arm. 

We can first use it to do a  **balance check** for Water Guard use across Arms.

```python
ever_yesno = WGP_ever.pivot('c2a_wg_used_ever','treatment_arm')
ever_yesno
```

Converting to percentages...

```python
total = ever_yesno.column(1) + ever_yesno.column(2)
ever_yesno = ever_yesno.with_columns('Percent No',ever_yesno.column(1) / total * 100, 
                                     'Percent Yes', ever_yesno.column(2) / total * 100)
ever_yesno
```

Let's also repeat the process for the variable of whether the households are currently using Water Guard, `c3a_wg_water_currently_treat`.

```python
WGP_current = WGP_table.where('c3a_wg_water_currently_treat',are.not_equal_to(0))
new2 = translate_to_yesno(WGP_current,'c3a_wg_water_currently_treat')
WGP_current = WGP_current.with_column('c3a_wg_water_currently_treat',new2)
WGP_current.group("c3a_wg_water_currently_treat")
```

Do you notice a problem here? Look at the total numbers reported in the output above.

We can do the same percentage tables for the balance check but maybe there's a problem. 
Look at the total number of households answering the question and compare that to the total number from the previous section.

```python
current_yesno = WGP_current.pivot('c3a_wg_water_currently_treat','treatment_arm')
total = current_yesno.column(1) + current_yesno.column(2)
current_yesno = current_yesno.with_columns('Percent No',current_yesno.column(1)/total * 100, 
                                           'Percent Yes', current_yesno.column(2)/total * 100)
current_yesno
```

This seems like a really high usage, but **maybe this is due to missing values**. 

Let's now also include the 0 (missing) values in our analysis.


```python
current_yesnomissing = WGP_table.pivot('c3a_wg_water_currently_treat','treatment_arm')
total = current_yesnomissing.column(1) + current_yesnomissing.column(2) + current_yesnomissing.column(3)
current_yesnomissing = current_yesnomissing.with_columns(
                                     'Percent Missing',current_yesnomissing.column("0.0") / total * 100, 
                                     'Percent No',current_yesnomissing.column("2.0") / total * 100, 
                                     'Percent Yes', current_yesnomissing.column("1.0") / total * 100)
current_yesnomissing
```

**Non-Graded Reflection Question:**


- Explain the previous table clearly and concisely, as if you were explaining it to someone who didn't know the back story.
- What does each rows/column mean?
- Conduct a balance check: does the distribution of yes/no/missing look balanced across arms?


<div id="Baseline"></div>

## Baseline and a Randomly Selected Compound




Let's describe a household selected at random.

First, we will extract the household/compound id into an array.

```python
hhld_array = WGP_table.column('a1_cmpd_id')
hhld_array
```

Next, we will draw randomly from this array.

```python
random_generator = np.random.default_rng(seed=10) # Please keep this line so the code is deterministic.
randomhh = random_generator.choice(hhld_array)
print("My randomly selected household is household number", randomhh)
```

Then, let's look at the data for our randomly selected household:

```python
myfamily = WGP_table.where("a1_cmpd_id",randomhh)
myfamily
```

Some of the variables may need some manipulation. 
Let's start with the age of the respondent:

```python
birthyear = myfamily.column("b3_birth_year").item(0)
surveyyear = myfamily.column("a5_date_interview_year").item(0)
agecalc = surveyyear-birthyear  # 
agecalc
```

And their tribe:

```python
print("Survey respondent Tribe", myfamily.column("b5_tribe").item(0))
print("Respondent Spouse Tribe", myfamily.column("b7_tribe_spouse").item(0))
```

Lastly, whether they have a latrine:

```python
print("Does the household have a latrine?", myfamily.column("d3_latrine").item(0))
```

Remember in the answer above it is coded so that 1=Yes and 2=No.

**Non-Graded Reflection Question:**
 Describe your randomly selected household and the respondent who is answering the survey. Please remember you can find the code sheet under the section of Baseline Survey.

1. Age
2. Tribe (Ethnicity)
3. Education 
4. Member of any groups b11-b15?
5. Occupation
6. Religion 
7. A summary of D variables, iron roof, floor materials, latrine, cattle, and others
8. Have they ever used WG?
9. Their treatment arm assignment
10. How many children do they have  
11. Gender and Age of children
12. Have any of the children been sick?


<div id="Chlorine"></div>

## Water Guard Usage outcome variables




### WGP Followup - Variability
The purpose of this section will be to continue on with the follow-up rounds of the Water Guard Promotion study.   In this section we have both the household reported use, and the use validated by checking the chlorine content of the water using a test kit.  

```python
WGP3rds_table = Table.read_table('WGP_3waves_Data8.csv')
WGP3rds_table
```

This is a large dataset, basically three datasets merged together, one for baseline, one for short term follow up and one for long term followup. The column `round` describes these 3 time steps:

- Round = 1 : baseline
- Round = 2 : 3 week followup
- Round = 3 : 3 month followup

Notably, many of the variables are only asked in one of the three rounds. For example, the chlorine use variables are:

- The variable for self reported chlorine use was `c6n` in Round 2, and `c5n` in Round 3.
- The variable for chlorine use is `c12n21pnk` in Round 2 and `c15npt2or1pnk` in Round 3.

Instead, the following variables have been combined across rounds for the ease of programming:

- `Selfrptpct` is self reported chlorine use in both round 2 and round 3
- `Vldclpct`  is validated chlorine use in both rounds

```python
WGP3rds_table.group("treatment_arm")
```

```python
WGP3rds_table.group('round')
```

### Grouping by round + treatment arm

We want to create a multi-level group: each group should be a unique combination of the survey round and the treatment arm.  


```python
WGP_3rds_outcomesonly= WGP3rds_table.select("round", "treatment_arm", "Selfrptpct", "Vldclpct")
WGP_3rds_outcomesonly.group(["round","treatment_arm"], np.mean).show(30)
```

### Making a smaller dataset

Lets break out a smaller dataset of the variables we want to focus on; just for Round 2 and the outcome variables.

```python
WGPRd2 = WGP3rds_table.where("round", 2).select("a1_cmpd_id","treatment_arm",
                                           "c6_current_water_treated_wg", 
                                           'c6_curr_water_treat_other_c',
                                           'c12_chlorine_meter_reading',
                                           'c11_chlorine_color','c12n21pnk', 'c6n'
                                          )
WGPRd2
```

A quick examination of the estimated Water Guard usage in Round 2 across all treatment arms:

```python
np.mean(WGPRd2.column('c12n21pnk'))
```

### A/B Testing

To see if the treatment was significant, we can utilize A/B testing. Recall from Data 8 that an A/B tests if two numerical samples come from the same underlying distribution, simulating the null hypothesis by shuffling the treatment assignment labels. 

Let's conduct an A/B testing to compare the difference in chlorine use (using the measures validated by a Chlorine measurement) between treatment and control. For this exercise, we will compare between arm 1 (the control) and arm 5 (flat-fee promoter + coupon).

First, let's calculate the observed difference between the two groups.


```python
relevant_households = WGPRd2.where("treatment_arm", are.contained_in(make_array(1, 5)))

grouped_tbl = relevant_households.group("treatment_arm", np.mean)
obs_diff = grouped_tbl.column("c12n21pnk mean").item(1) - grouped_tbl.column("c12n21pnk mean").item(0)
print("The observed difference is:", obs_diff )
```

Now we will simulate under the null, shuffling the treatment assignment labels.

```python
mean_diffs = make_array()
for i in np.arange(1000):
    shuffled_labels = relevant_households.sample(with_replacement = False).column("treatment_arm")
    shuffled_tbl = relevant_households.with_column("shuffled treatment", shuffled_labels)
    grouped_tbl = shuffled_tbl.group("shuffled treatment", np.mean)
    mean_diff = grouped_tbl.column("c12n21pnk mean").item(1) - grouped_tbl.column("c12n21pnk mean").item(0)
    mean_diffs = np.append(mean_diffs, mean_diff)
Table().with_columns("Difference in means", mean_diffs).hist()
plt.scatter(obs_diff, 0, c = 'red')
```

### Computing Confidence Intervals

Another way to compare between two groups is to examine and compare confidence intervals of each group's mean.
Let's compute a confidence interval for the percent of households using Chlorine (the measures validated by a Chlorine measurement) via resampling.


```python
chlorine_uses = make_array()
for i in np.arange(1000):
    bootstrapped_sample = WGPRd2.sample()
    sample_mean = np.mean(bootstrapped_sample.column('c12n21pnk'))
    chlorine_uses = np.append(chlorine_uses, sample_mean)
Table().with_columns("Chlorine use mean", chlorine_uses).hist()
```

```python
lower = percentile(2.5, chlorine_uses)
upper = percentile(97.5, chlorine_uses)
print(f"The 95% confidence interval is [{lower}, {upper}]")
```

Using this technique we can look at the Confidence Intervals for each of the treatment arms.

```python
# Arm 1
arm_1 = WGPRd2.where("treatment_arm", 1)
chlorine_uses_arm1 = make_array()
for i in np.arange(1000):
    bootstrapped_sample = arm_1.sample()
    sample_mean = np.mean(bootstrapped_sample.column('c12n21pnk'))
    chlorine_uses_arm1 = np.append(chlorine_uses_arm1, sample_mean)
Table().with_columns("Chlorine use mean (arm 1)", chlorine_uses_arm1).hist()
```

```python
lower = percentile(2.5, chlorine_uses_arm1)
upper = percentile(97.5, chlorine_uses_arm1)
print(f"The 95% confidence interval for arm 1 is [{lower}, {upper}]")
```

Repeating for Arm 2...


```python
arm_2 = WGPRd2.where("treatment_arm", 2)
chlorine_uses_arm2 = make_array()
for i in np.arange(1000):
    bootstrapped_sample = arm_2.sample()
    sample_mean = np.mean(bootstrapped_sample.column('c12n21pnk'))
    chlorine_uses_arm2 = np.append(chlorine_uses_arm2, sample_mean)
Table().with_columns("Chlorine use mean (arm 2)", chlorine_uses_arm2).hist()
```

```python
lower = percentile(2.5, chlorine_uses_arm2)
upper = percentile(97.5, chlorine_uses_arm2)
print(f"The 95% confidence interval for arm 2 is [{lower}, {upper}]")
```

**Question 2.1:**
What can we tell by comparing the confidence intervals for Arm 1 and Arm 2? Do they overlap?  What does that mean?

<div>
a. The confidence intervals for Arm 1 and Arm 2 overlap. This means that they are not significantly different from each other.
<br>
b. The confidence intervals for Arm 1 and Arm 2 do not overlap. This means that they are not significantly different from each other.
<br>
c. The confidence intervals for Arm 1 and Arm 2 overlap. This means that they are significantly different from each other.
<br>
d. The confidence intervals for Arm 1 and Arm 2 do not overlap. This means that they are significantly different from each other.
<br>
</div>

Assign the letter corresponding to your answer to `q2_1` below.


```python
q2_1 = ...
```

```python
grader.check("q2_1")
```

<div id="Graph"></div>

## Graph of outcomes by Treatment Arm




### WGP Followup Round 2, Round 3, and across rounds

Now we will work on making a summary graph for the WGP study. This graph should show the seven treatment arms with the levels of our outcomes, and include error bars (the standard deviations of our sample means). In addition, we can work on customizing and saving our graph.  

```python
WGPRd2 = WGP3rds_table.where("round", 2).select("a1_cmpd_id",'treatment_arm','Selfrptpct', 'Vldclpct')
WGPRd2
```

```python
# Group by treatment_arm and take the means of each group
round2_means = WGPRd2.group('treatment_arm', np.mean)
round2_means
```

```python
# Save the means into an array for later use
round2_means_self_array = round2_means.column('Selfrptpct mean')
round2_means_vld_array = round2_means.column('Vldclpct mean')
```

Let's start with a bar chart of the self reported Water Guard usage across treatment arms.

```python
round2_means.bar('treatment_arm','Selfrptpct mean') 
```

Comparing the self reported values against the validated values:

```python
round2_means.bar('treatment_arm',make_array(2, 3)) 
```

Next, let's redo the same procedure for round 3.

```python
WGPRd3 = WGP3rds_table.where("round",3).select("a1_cmpd_id",'treatment_arm','Selfrptpct', 'Vldclpct')
round3_means = WGPRd3.group('treatment_arm', np.mean)
round3_means_array = round3_means.column('Selfrptpct mean')
round3_means.bar('treatment_arm' ,make_array(2, 3)) 
```

Congrats! You've finished Lab 6!

---
**References:**
1. Rao & Tam curve https://www.tandfonline.com/doi/ref/10.1080/02664768700000032
2. Income and Poverty in the US https://www.census.gov/library/publications/2019/demo/p60-266.html

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> assert len(cum_income_share) == 5
```
```
>>> assert np.sum(np.diff(cum_income_share) >= 0) == 4
```
```
>>> assert np.isclose(sum(us_cumulative.column("income_percentile")), 3.0)
```
```
>>> assert np.isclose(sum(us_cumulative.column("cum_income_share")), 1.881)
```


Tests for q1_2 (points: None)
```
>>> assert type(x_val) == float
```
```
>>> assert type(y_val) == float
```
```
>>> assert x_val > 0.5 and y_val > 0.5
```


Tests for q1_3 (points: None)
```
>>> assert us_cumulative_extended.num_rows == 6
```
```
>>> assert sum([np.isclose(i, .95) for i in us_cumulative_extended["income_percentile"]]) == 1
```
```
>>> assert sum([np.isclose(i, .769) for i in us_cumulative_extended["cum_income_share"]]) == 1
```


Tests for q1_4 (points: [0])
```
>>> assert q1_4 in ["a", "b", "c", "d"]
```


Tests for q1_5 (points: None)
```
>>> assert us_log.num_rows == 6
```
```
>>> assert us_log.num_columns == 2
```
```
>>> assert sum(us_log.column("log_income_percentile")) < 0
```


Tests for q2_1 (points: [0])
```
>>> assert q2_1 in ["a", "b", "c", "d"]
```
