```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

<div style="display: flex; justify-content: space-between; align-items: center; background-color: transparent;">
    <div>
        <img src="https://data-88e.github.io/assets/images/blue_text.png" width="250px" style="margin-left: 0;" />
    </div>
    <div style="text-align: right; font-size: 10pt;">
        <strong>Economic Models</strong>, edX<br>
        Dr. Eric Van Dusen<br>
        Shashank Dalmia<br>
        Amal Bhatnagar<br>
        Akhil Venkatesh
    </div>
</div>

## Lab 2: Supply

**Note:** For this lab, two line answers are more than enough for full credit on written answers!

```python
from datascience import  *
import matplotlib.pyplot as plt
%matplotlib inline
import numpy as np
import sympy

def plot_equation(equation, price_start, price_end, label=None):
    plot_prices = [price_start, price_end]
    plot_quantities = [equation.subs(list(equation.free_symbols)[0], c) for c in plot_prices]
    plt.plot(plot_quantities, plot_prices, label=label)
    
def plot_intercept(eq1, eq2):
    ex = sympy.solve(eq1-eq2)[0]
    why = eq1.subs(list(eq1.free_symbols)[0], ex)
    plt.scatter([why], [ex], zorder=10, color="tab:orange")
    return (ex, why)

solve = lambda x,y: sympy.solve(x-y)[0] if len(sympy.solve(x-y))==1 else "Not Single Solution"
```

### Question 1 

Fill in the blank entries in the table below. Note that the output values are not linearly increasing. Two rows for values at levels 8 and 9 are intentionally left missing.

| Output | Total Fixed Cost | Total Variable Cost | Total Cost | Average Fixed Cost | Average Variable Cost | Average Total Cost | Marginal Cost |
|--------|------------------|---------------------|------------|--------------------|-----------------------|--------------------|---------------|
| 0 | 60 | **(i)** | 60 | 0 | 0 | 0 | 0 |
| 1 | 60 | 40 | 100 | 60 | 40 | 100 | **(ii)** |
| 2 | 60 | 65 | 125 | 30 | 33 | 63 | 25 |
| 3 | 60 | 85 | 145 | 20 | 28 | 48 | 20 |
| 4 | 60 | 100 | 160 | 15 | 25 | **(iii)** | 15 |
| 5 | 60 | 120 | 180 | 12 | 24 | 36 | 20 |
| 6 | 60 | 144 | 204 | 10 | 24 | 34 | 24 |
| 7 | 60 | 182 | 242 | 9 | 26 | 35 | 38 |
| 8 | 60 | 240 | 300 | 8 | **(iv)** | 38 | 58 |
| 10 | 60 | 410 | 470 | 6 | 41 | 47 | 85 |
| 15 | 60 | **(v)** | 970 | 4 | 60.67 | 64.67 | **(vi)** |
| 25 | 60 | **(vii)** | **(viii)** | 2 | NA | NA | 125 |


```python
i = ...
```

```python
grader.check("q1_1")
```

```python
ii = ...
```

```python
grader.check("q1_2")
```

```python
iii  = ...
```

```python
grader.check("q1_3")
```

```python
iv = ...
```

```python
grader.check("q1_4")
```

```python
v = ...
```

```python
grader.check("q1_5")
```

```python
vi = ...
```

```python
grader.check("q1_6")
```

```python
vii = ...
```

```python
grader.check("q1_7")
```

```python
viii = ...
```

```python
grader.check("q1_8")
```

### Question 2

Use the table from Question 1 to answer the following questions:

If the price of the good is $50, will the firm produce? If so, how much? Assign `True` or `False` to `firms_will_produce` and the amount that the firm will produce to `output`. If you believe the firm will not produce, assign `0` to `output`.

*Hint*: Consider the marginal costs at each amount of production. 


```python
firm_will_produce = ...
output = ...
```

```python
grader.check("q2_1")
```

### Question 3

Referring to the last plot titled 'All Energy Sources' (the colorful bar plot) in this week's lecture, is it possible to calculate the profit of any firm at any price with the given data (if they produce)? Assign `"Yes"` or `"No"` to `q3`. Your response should only be considering the information in the bar plot. 


```python
q3 = ...
```

```python
grader.check("q3")
```

### Question 4
(Optional, Non-Graded Reflection Question)
Why does the Marginal Cost line intersect with the Average Variable Cost and Average Total Cost at their minima?


### Question 5

What happens to the Average Fixed Cost (AFC) and the difference between Average Total Cost (ATC) and Average Variable Cost (AVC) as the quantity increases? Assign `decreases`, `increases` or `stays same` to the respective results.

```python
AFC = ...
ATC_minus_AVC = ...
```

```python
grader.check("q5")
```

###  Question 6

We will return to our exploration of beef from lab 1 to solve for market equilibria, using SymPy to create symbolic expressions for the supply and demand curves. Let's start out by creating an upward sloping supply curve, where $Q_S$ is quantity supplied and $P_S$ is the price at that supplied quantity. 


```python
beef = Table.read_table('beef.csv')
beef
```

**Question 6.1**: Create symbols for the variables $P_S$ and $Q_S$ and assign them to `P_S` and `Q_S`, respectively.

_Hint:_ To make a subscript in SymPy, we can use LaTeX syntax; for example, $a_b$ would be created through `sympy.Symbol("a_b")`.


```python
P_S = ...
Q_S = ...
```

```python
grader.check("q6_1")
```

**Question 6.2:** Determine the linear demand curve (not a log demand curve), expressed using quantity as a function of price, by using `np.polyfit`. Assign the results to `demand_curve_properties`. 


```python
demand_curve_properties = ...
demand_curve_properties
```

```python
grader.check("q6_2")
```

**Question 6.3:**
Create a new symbol $P$ for the demand price and assign it to `P`. Then construct the demand curve from the `np.polyfit` results in the previous part using SymPy and `P`. Make sure to use the values directly from `demand_curve_properties` so that you don't round values, or else the autograder may fail.

Recall that we can create a function of symbol `x` by using normal math operations: 

```python
x = sympy.Symbol("x")
y = 2 * x + 1
```


```python
P = ...
demand = demand_curve_properties.item(...) * ... + demand_curve_properties.item(...)
demand
```

```python
grader.check("q6_3")
```

**Question 6.4:**
Interpret the demand curve. What do its slope and intercept mean? Assume you have correct `demand_slope` and `demand_intercept` from the previous question.

1. The slope tells us that a 1 dollar decrease in price results in `demand_slope` less beef demanded. The intercept states that at `demand_intercept` dollars, the total number of beef demanded is 0.
2. The slope tells us that a 1 dollar increase in price results in `demand_slope` less beef demanded. The intercept states that at 0 dollars, the total number of beef demanded is `demand_intercept`.
3. The slope tells us that a 1 dollar decrease in price results in `demand_slope` less beef demanded. The intercept states that at 0 dollars, the total number of beef demanded is `demand_intercept`.
4. The slope tells us that a 1 dollar increase in price results in `demand_slope` less beef demanded. The intercept states that at `demand_intercept` dollars, the total number of beef demanded is 0. 

Assign the number corresponding to your answer to `q6_4` below.


```python
q6_4 = ...
```

```python
grader.check("q6_4")
```

**Question 6.5:** 
As we will learn, the supply curve is the relationship between the price of a good or service and the quantity of that good or service that the seller is willing to supply. Let's assume the supply curve for beef is  $S(P)=0.094P$. This supply curve is not based on any data. Use SymPy to create an expression for supply and assign this to the variable `supply`.

```python
supply = ...
supply
```

```python
grader.check("q6_5")
```

**Question 6.6:**
Interpret the supply curve. What do its slope and intercept mean? How does supply change with the price of beef?

1. As price decreases, the supply of beef increases. For every dollar decrease, we expect 0.094 increase in supply. At \$0, there is no supply (0 units supplied).
2. As price increases, the supply of beef decrease. For every dollar increase, we expect 0.094 decrease in supply. At \$0, there are 0.094 units supplied.
3. As price increases, the supply of beef increases. For every dollar increase, we expect 0.094 increase in supply. At \$0, there is no supply (0 units supplied).
4. As price increases, the supply of beef increases. For every dollar increase, we expect 0.094 increase in supply. At \$0, there are 0.094 units supplied.

Assign the number corresponding to your answer to `q6_6` below.


```python
q6_6 = ...
```

```python
grader.check("q6_6")
```

**Question 6.7:**
The price equilibrium is the price at which the supply curve and demand curve intersect. In this case, the price equilibrium would be both the price of beef that consumers want to purchase at as well as the price of beef that the manufacturers want to sell at. Find the price equilibrium $P^*$ using the supply and demand curves for each market. Assign this value to `P_star`.

_Hint:_ Our `solve` function may be helpful here...


```python
P_star = ...
P_star
```

```python
grader.check("q6_7")
```

**Question 6.8:** 
The quantity equilibrium is the quantity at which the supply curve and demand curve intersect. In this case, the quantity equilibrium would be both the quantity of beef that consumers want to purchase as well as the amount of beef manufacturers want to provide. Find the quantity equilibrium $Q^*$ using the supply and demand curves for each market. Assign this value to `Q_star`.


```python
Q_star = ...
Q_star
```

```python
grader.check("q6_8")
```

**Question 6.9:** 
Assume that due to favorable economic conditions, consumers now have a higher disposable income and are buying more beef. How will this affect market equilibrium?

1. Increase in price equilbrium due to the income effect.
2. Decrease in price equilbrium due to the income effect.
3. Increase in price equilbrium due to the substitution effect.
4. Decrease in price equilbrium due to the substitution effect.

Assign the number corresponding to your answer to `q6_9` below.


```python
q6_9 = ...
```

```python
grader.check("q6_9")
```

**Question 6.10:** 
Now assume that there is an avian flu outbreak and the price of chickens goes through the roof (increases severely)! How will this affect market equilibrium for the beef market?

1. Increase in price equilbrium due to the income effect.
2. Decrease in price equilbrium due to the income effect.
3. Increase in price equilbrium due to the substitution effect.
4. Decrease in price equilbrium due to the substitution effect.

Assign the number corresponding to your answer to `q6_10` below.


```python
q6_10 = ...
```

```python
grader.check("q6_10")
```

**Congratulations! You just finished Lab 2!**
<br>

Tests for q1_1 (points: 1)
```
>>> type(i) == int
True
```


Tests for q1_2 (points: 1)
```
>>> type(ii) == int
True
```


Tests for q1_3 (points: 1)
```
>>> type(iii) == int
True
```


Tests for q1_4 (points: 1)
```
>>> type(iv) == int
True
```


Tests for q1_5 (points: 1)
```
>>> type(v) == int
True
```


Tests for q1_6 (points: 1)
```
>>> type(vi) == int
True
```


Tests for q1_7 (points: 1)
```
>>> type(vii) == int
True
```


Tests for q1_8 (points: 1)
```
>>> type(viii) == int
True
```


Tests for q2_1 (points: [0])
```
>>> assert firm_will_produce in [True, False]
>>> assert output >= 0
```


Tests for q3 (points: [0])
```
>>> assert q3 in ["Yes", "No"]
```


Tests for q5 (points: [0])
```
>>> assert AFC in ["decreases", "increases", "stays same"] and ATC_minus_AVC in ["decreases", "increases", "stays same"]
```


Tests for q6_1 (points: None)
```
>>> str(P_S) == "P_S"
True
```
```
>>> str(Q_S) == "Q_S"
True
```
```
>>> type(P_S) == sympy.Symbol
True
```
```
>>> type(Q_S) == sympy.Symbol
True
```


Tests for q6_10 (points: [0])
```
>>> assert q6_10 in [1, 2, 3, 4]
```


Tests for q6_2 (points: None)
```
>>> len(demand_curve_properties) == 2
True
```


Tests for q6_3 (points: None)
```
>>> assert type(P) == sympy.Symbol
>>> assert str(P) == "P"
```
```
>>> assert len(demand.as_coefficients_dict()) == 2 # check that it is a line
>>> assert P in demand.as_coefficients_dict() # check that the line has Q as its variable
>>> assert 1 in demand.as_coefficients_dict() # check that the line has an intercept
```


Tests for q6_4 (points: [0])
```
>>> assert q6_4 in [1, 2, 3, 4]
```


Tests for q6_5 (points: None)
```
>>> assert P in supply.as_coefficients_dict()
```


Tests for q6_6 (points: [0])
```
>>> assert q6_6 in [1, 2, 3, 4]
```


Tests for q6_7 (points: None)
```
>>> 0 <= P_star <= 1000
True
```


Tests for q6_8 (points: None)
```
>>> 15 <= Q_star <= 50
True
```


Tests for q6_9 (points: [0])
```
>>> assert q6_9 in [1, 2, 3, 4]
```
