```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

<div id="nb-header" style="display: flex; justify-content: space-between; align-items: center; background-color: transparent;">
    <div>
        <img src="https://data-88e.github.io/assets/images/blue_text.png" width="250px" style="margin-left: 0;" />
    </div>
    <div style="text-align: right; font-size: 10pt;">
        <strong>Economic Models</strong>, edX<br>
        Dr. Eric Van Dusen <br>
        Amal Bhatnagar <br>
        Alan Liang <br>
        Akhil Venkatesh <br>
    </div>
</div>

# Lab 1: Demand

```python
import numpy as np
from datascience import *
import sympy

import matplotlib as mpl
import matplotlib.pyplot as plt
%matplotlib inline
mpl.rcParams['figure.dpi'] = 100

import warnings
warnings.filterwarnings('ignore')

solve = lambda x,y: sympy.solve(x-y)[0] if len(sympy.solve(x-y))==1 else "Not Single Solution"

def plot_equation(equation, price_start, price_end, label=None):
    plot_prices = [price_start, price_end]
    plot_quantities = [equation.subs(list(equation.free_symbols)[0], c) for c in plot_prices]
    plt.plot(plot_prices, plot_quantities, label=label)
    
def plot_intercept(eq1, eq2):
    ex = sympy.solve(eq1-eq2)[0]
    why = eq1.subs(list(eq1.free_symbols)[0], ex)
    plt.scatter([ex], [why])
    return (ex, why)
```

## Demand Curve for Beef

In today's lab, we will explore the demand curve for beef across time. The data is originally from the US Department of Agriculture and examines how the price varies between 1977 and 1999. 

For the purposes of this assignment, the quantity of beef demanded is measured in millions of pounds, while price is measured in average cents per pound. 

Note that throughout this exercise, we will ignore the effects of inflation. 

Remember, if you're feeling stuck on a problem, please refer to [these pages]( https://data-88e.github.io/textbook/content/01-demand/02-example.html)!

```python
beef = Table.read_table('beef.csv')
beef
```

To construct the demand curve, let's first examine what the relationship between price and quantity is. We would expect to see a downward-sloping line between price and quantity; if a product's price increases, consumers will purchase less of it, and if a product's price decreases, then consumers will purchase more of it.

<!-- BEGIN QUESTION -->

**Question 1:** 
Create a scatterplot with a best-fit line (using `fit_line = True`) across the points. Trend-lines are helpful because they consolidate all the datapoints into a single line, helping us better understand the relationship between the two variables. 


```python
...
plt.title("Demand Curve for Beef");
```

<!-- END QUESTION -->

We will now quantify our demand curve using the [`np.polyfit(...)`](https://numpy.org/doc/stable/reference/generated/numpy.polyfit.html) function, which returns the coefficients of a fitted $d$-degree polynomial. 

`np.polyfit` takes 3 parameters
- array of x-coordinates
- array of y-coordinates
- degree of polynomial 

Because we are looking for a **linear** function to serve as the demand curve, we will use 1 as the degree of polynomial. It will therefore return an array of size 2, where the first element is the `slope` and the second is the `y-intercept`.

**Question 2.1:** Create a non-log demand curve that expresses the price of beef in terms of the quantity demanded (using quantity as $x$ and price as $y$) by using `np.polyfit`. Assign the results to `non_log_demand_params_1`. 



```python
non_log_demand_params_1 = ...
non_log_demand_params_1
```

```python
grader.check("q2_1")
```

Your answer should show that the $m$ variable (slope) is -19.37 and b variable (y-intercept) is 607. Using the $y = mx + b$ template above, then our demand curve would be $y = -19.37x + 607$

**Question 2.2:** Assign the number (i.e `1`, `2`, `3`, or `4`) to `slope_intercept` that best represents the interpretation of the slope and intercept terms of the demand curve in question `2.1`.

*Hint*: watch [this video](https://www.youtube.com/watch?v=LwLh6ax0zTE)

1. For a decrease in 19.37 million pounds of beef demanded per quarter, we would expect to see a decrease in the price per pound of beef by 1 cent. At a price of 607 cents per pound, we would expect there to be 0 pounds of beef demanded.
2. For an increase in 1 million pounds of beef demanded per quarter, we would expect to see a decrease in the price per pound of beef by 19.37 cents. At a price of 607 cents per pound, we would expect there to be 0 pounds of beef demanded.
3. For an increase in 1 million pounds of beef demanded per quarter, we would expect to see a decrease in the price per pound of beef by 19.37 cents. At a price of 0 cents per pound, we would expect there to be 0 pounds of beef demanded.
4. For an increase in 1 million pounds of beef demanded per quarter, we would expect to see an increase in the price per pound of beef by 19.37 cents. At a price of 607 cents per pound, we would expect there to be 0 pounds of beef demanded.

```python
slope_intercept = ...
```

```python
grader.check("q2_2")
```

Let's visualize our non-log demand curve from question 2:

```python
quantities = np.arange(15,25)
predicted_prices_non_log = non_log_demand_params_1[0] * quantities + non_log_demand_params_1[1]

plt.figure(figsize=(8,6))
plt.scatter(beef.column("Quantity"), beef.column("Price"))
plt.plot(quantities, predicted_prices_non_log, color = 'red', label = 'Non-log demand curve (P as function of Q)')
plt.xlabel('Quantity')
plt.ylabel('Price')
plt.legend();
```

Our interpretation of the demand curve expressed as price in terms of quantity above was not particularly intuitive: changes in quantity demanded likely do not trigger changes in price, but instead the other way around. Instead, we can determine a new demand curve that expresses quantity as a function of price.



**Question 3.1:** Create a non-log demand curve that expresses the quantity of beef in terms of the price by using `np.polyfit`. Assign the results to `non_log_demand_params_2`. 



```python
non_log_demand_params_2 = ...
non_log_demand_params_2
```

```python
grader.check("q3_1")
```

**Question 3.2:** Let's say your friend Ergun is trying to find the demand of beef when the average price is 250 cents per pound. Use `non_log_demand_params_2` to find the exact demand, in millions of pounds of beef.


```python
demand_q3_2 = ...
```

```python
grader.check("q3_2")
```

Let's visualize our non-log demand curve from question 3:

```python
prices = np.arange(150,300)
predicted_quantities_non_log = non_log_demand_params_2[0] * prices + non_log_demand_params_2[1]

plt.figure(figsize=(8,6))
plt.scatter(beef.column("Quantity"), beef.column("Price"))
plt.plot(predicted_quantities_non_log, prices, color = 'red', label = 'Non-log demand curve (Q as function of P)')
plt.xlabel('Quantity')
plt.ylabel('Price')
plt.legend();
```

Now we will try fitting semi-log and log-log demand curves on the same data to see which model works the best. 

**Question 4.1:** Create a semi-log demand curve that expresses the quantity of beef in terms of the price by using `np.polyfit`. Assign the results to `semi_log_demand_params`. 



```python
log_quantity = ...

semi_log_demand_params = ...
semi_log_demand_params
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

**Question 4.2:** Assign the number (i.e `1`, `2`, or `3`) to `slope_only` that best represents the interpretation of the slope term of the demand curve in question `4.1`.

1. A 1 cent decrease in the price per pound of beef is associated with a 0.244% decrease in the quantity demanded per quarter.
2. A 1 cent increase in the price per pound of beef is associated with a 0.244% decrease in the quantity demanded per quarter.
3. A 1 cent increase in the price per pound of beef is associated with a 0.244% increase in the quantity demanded per quarter.



```python
slope_only = ...
```

```python
grader.check("q4_2")
```

<!-- END QUESTION -->

Let's visualize our semi-log demand curve:

```python
prices = np.arange(150,300)
predicted_quantities_semi_log = (np.e ** (prices * semi_log_demand_params[0])) * (np.e ** semi_log_demand_params[1])

plt.figure(figsize=(8,6))
plt.scatter(beef.column("Quantity"), beef.column("Price"))
plt.plot(predicted_quantities_semi_log, prices, color = 'red', label = 'Semi-log demand curve')
plt.xlabel('Quantity')
plt.ylabel('Price')
plt.legend();
```

**Question 5.1:** Create a log-log demand curve that expresses the quantity of beef in terms of the price by using `np.polyfit`. Assign the results to `log_log_demand_params`. 



```python
log_quantity = ...
log_price = ...

log_log_demand_params = ...
log_log_demand_params
```

```python
grader.check("q5_1")
```

<!-- BEGIN QUESTION -->

**Question 5.2:** Let's say that Matt wants to find the what a 15% increase in the price of beef is associated with in terms of a percentage change in the quantity demanded per quarter. Use `log_log_demand_params` to find the exact result and assign the result to `q5_2`.


```python
q5_2 = ...
```

```python
grader.check("q5_2")
```

<!-- END QUESTION -->

Let's visualize our log-log demand curve:

```python
prices = np.arange(150,300)
predicted_quantities_log_log = (prices ** log_log_demand_params[0]) * (np.e ** log_log_demand_params[1])

plt.figure(figsize=(8,6))
plt.scatter(beef.column("Quantity"), beef.column("Price"))
plt.plot(predicted_quantities_log_log, prices, color = 'red', label = 'Log-log demand curve')
plt.xlabel('Quantity')
plt.ylabel('Price')
plt.legend();
```

**Question 5.3:** Is the demand for beef...
1. perfectly inelastic
2. inelastic
3. neither elastic nor inelastic
4. elastic
5. perfectly elastic

Assign your answer as a number from 1-5 to `demand_elasticity`.



```python
demand_elasticity = ...
```

```python
grader.check("q5_3")
```

Let's visualize all 4 demand curves we've created:

```python
prices = np.arange(150,300)
predicted_quantities_log_log = (prices ** log_log_demand_params[0]) * (np.e ** log_log_demand_params[1])

plt.figure(figsize=(8,6))
plt.scatter(beef.column("Quantity"), beef.column("Price"), alpha = 0.5)
plt.plot(quantities, predicted_prices_non_log, color = 'purple', label = 'Non-log demand curve (P as function of Q)')
plt.plot(predicted_quantities_non_log, prices, color = 'orange', label = 'Non-log demand curve (Q as function of P)')
plt.plot(predicted_quantities_semi_log, prices, color = 'red', label = 'Semi-log demand curve')
plt.plot(predicted_quantities_log_log, prices, color = 'green', label = 'Log-log demand curve')
plt.xlabel('Quantity')
plt.ylabel('Price')
plt.legend();
```

**Question 6:** For each of the following two supply curves, assign one as `inelastic` and the other as `elastic`. 

```python
plt.figure(figsize = [6,6])
plt.plot(np.arange(0,10,1), np.flip(np.array([0,3,6,9,12,15,18,21,24,27])), label = 'curve 1')
plt.plot(np.flip(np.array([0,3,6,9,12,15,18,21,24,27])), np.arange(0,10,1), color = 'red', label = 'curve 2')
plt.xlim(0,30)
plt.ylim(0,30)
plt.xticks([], [])
plt.yticks([], [])
plt.legend();
```

```python
curve_1 = ...
curve_2 = ...
print("curve 1:", curve_1)
print("curve 2:", curve_2)
```

```python
grader.check("q6")
```

**Congratulations! You finished Lab 1!**

Tests for q2_1 (points: [0])
```
>>> len(non_log_demand_params_1) == 2
True
```


Tests for q2_2 (points: None)
```
>>> slope_intercept in [1,2,3,4,5]
True
```


Tests for q3_1 (points: [0])
```
>>> len(non_log_demand_params_2) == 2
True
```


Tests for q3_2 (points: [0])
```
>>> assert demand_q3_2 >= 0
```


Tests for q4_1 (points: [0])
```
>>> len(semi_log_demand_params) == 2
True
```


Tests for q4_2 (points: None)
```
>>> slope_only in [1,2,3]
True
```


Tests for q5_1 (points: [0])
```
>>> len(log_log_demand_params) == 2
True
```


Tests for q5_2 (points: None)
```
>>> assert q5_2 <= 0
```


Tests for q5_3 (points: None)
```
>>> demand_elasticity in [1,2,3,4,5]
True
```


Tests for q6 (points: None)
```
>>> curve_1 in ["elastic", "inelastic"]
True
```
```
>>> curve_2 in ["elastic", "inelastic"]
True
```
