```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

<div id="nb-header" style="display: flex; justify-content: space-between; align-items: center; background-color: transparent;">
    <div>
        <img src="https://data-88e.github.io/assets/images/blue_text.png" width="250px" style="margin-left: 0;" />
    </div>
    <div style="text-align: right; font-size: 10pt;">
        <strong>Economic Models</strong>, edX<br>
        Dr. Eric Van Dusen<br>
        Alan Liang<br>
        Umar Maniku<br>
        Akhil Venkatesh
    </div>
</div>

<!-- BEGIN QUESTION -->

# Lab 3: Taxes, Government Intervention, and Welfare

<!-- END QUESTION -->

```python
from datascience import *
import numpy as np
import pandas as pd
import sympy
solve = lambda x,y: sympy.solve(x-y)[0] if len(sympy.solve(x-y))==1 else "Not Single Solution"
%matplotlib inline
```

## Question 1

Many states in the United States impose an excise and an ad valorem tax on the consumption of beer. According to the Beer Institute, “taxes are the single most expensive ingredient in beer, costing more than the labor and raw materials combined.” In fact, approximately 40 percent of the retail price of beer is dedicated toward covering all the applicable taxes. 

Throughout this exercise, you will investigate how salient consumers are to changes in taxes on beer. Typically, the ad valorem tax on beer is not reflected in the label price, like typical sales taxes in the US. However, the excise tax is included in the label price. 

This study on consumer salience towards taxes is taken from [_Salience and Taxation: Theory and Evidence_](https://www.aeaweb.org/articles?id=10.1257/aer.99.4.1145) by Chetty, Looney, and Kroft (AER 2010). 

**Question 1.1:**
In the paper *Salience and Taxation: Theory and Evidence*, the authors also investigate consumer salience towards taxes through a field experiment by posting tax-inclusive price tags, and find that doing so reduces demand by 8 percent. What does this imply about consumer salience in relation to the decrease in demand?


1. Including the tax in the price tag does not change the actual price of the good, but only the label price. Since demand decreases when including the tax in the price, we can conclude that not all consumers are aware of the tax.
2. Including the tax in the price tag does not change the actual price of the good, but only the label price. Since demand increases when including the tax in the price, we can conclude that not all consumers are aware of the tax.
3. Including the tax in the price tag changes the actual price of the good. Since demand decreases when including the tax in the price, we can conclude that all consumers are aware of the tax.
4. Including the tax in the price tag changes the actual price of the good, but only the signaling price. Since demand increases when including the tax in the price, we can conclude that not all consumers are aware of the tax.

```python
q1_1 = ...
```

```python
grader.check("q1_1")
```

Let's load in a table about beer sales across the 50 states from 1970 to 2003.

The table `beer_raw` includes 7 columns:
- `st_name`: the state abbreviation
- `year`: the year the data was recorded
- `c_beer`: the quantity of beer consumed in gallons (in thousands of gallons)
- `beer_tax`: the ad valorem tax as a percentage
- `btax_dollars`: the excise tax, represented in dollars per case of beer (24 cans)
- `population`: the population of the state (in thousands)
- `salestax`: the sales tax percentage

```python
beer_raw = Table.read_table("beer_tax.csv")
beer_raw
```

**Question 1.2:**
Consumption across states vary widely due to differences in population. Let's normalize this by adding a column called `pc_beer`, representing the per capita (per person) consumption of beer in each year for each state, to `beer_raw`. Call the resulting table `beer`.


```python
pc_beer_column = ...
beer = ...
beer
```

```python
grader.check("q1_2")
```

**Question 1.3:**
For the rest of this question, we will restrict our analysis to only one state - California. Assign `beer_state` to a filtered table with only data describing California.


```python
beer_state = ...
beer_state
```

```python
grader.check("q1_3")
```

**Question 1.4:**
To determine consumers' sensitivity to a change in taxes, we need to determine the changes in ad valorem taxes between each year and the corresponding change in the per capita beer consumption. Construct a new table `beer_state_changes` with these  2 columns: `tax_diff` and `consumption_diff`.

_Hint:_ Use `np.diff`.


```python
tax_diff = ...
per_capita_consumption_diff = ...
beer_state_changes = ...
beer_state_changes
```

```python
grader.check("q1_4")
```

**Question 1.5:**
Now we are ready to plot our results. What kind of graph will be most appropriate to denote the relationship between changes in taxes and changes in consumption?

<ol style="list-style-type: lower-alpha;">
    <li>Line plot</li>
    <li>Histogram</li>
    <li>Scatter plot</li>
    <li>Bar chart</li>
</ol>

Assign `q5` to letter corresponding to your choice.


```python
q1_5 = ...
```

```python
grader.check("q1_5")
```

<!-- BEGIN QUESTION -->

**Question 1.6:**
Plot the results from `beer_state_changes` using a scatter plot.


```python
...
```

<!-- END QUESTION -->



We will conduct a linear regression to determine the slope and intercept.

```python
beer_state_changes.scatter(0, 1, fit_line=True)
```

What are the values of the intercept and slope terms? 

To answer this, we can use `np.polyfit` with degree 1, which returns the slope and intercept terms in that order.

```python
np.polyfit(x = beer_state_changes.column("tax_diff"), y=beer_state_changes.column("consumption_diff"), deg=1)
```

Now let's see if the phenomenon applies to all states, across all years.



```python
# You do not have to understand how the code below works, but it works very similarly to what you've done above!
data = pd.read_stata("alcohol_data.dta")
data = data[["st_name", "year", "c_beer", "population", "beer_tax", "salestax"]]
data['beer_pc'] = data.c_beer / data.population
data = data.dropna(axis = 0)
def state_group(df):
    tax_diff = np.diff(df["beer_tax"])
    cons_diff = np.diff(df["beer_pc"])
    return pd.DataFrame({"tax_diff": tax_diff, "cons_diff": cons_diff})
differences = data.groupby("st_name").apply(state_group)
diff = Table.from_df(differences)
diff.scatter(0, 1)
```

The data seems to be a little messy, but a downward sloping trend seems to appear. Nonetheless, let's try to clean it up a little. 

We will aggregate the data based on state so that each state will be one data point. Each point represents a state's average change in taxes and consumption across the years.

```python
# You also do not have to understand how the code below works.
def state_group(df):
    tax_diff = np.mean(np.diff(df["beer_tax"]))
    cons_diff = np.mean(np.diff(df["beer_pc"]))
    return pd.DataFrame({"tax_diff": tax_diff, "cons_diff": cons_diff}, index=[0])
differences = data.groupby("st_name").apply(state_group)
diff = Table.from_df(differences)
diff.scatter(0, 1)
```

Up until now, we've been examining the effect of changes in the ad valorem tax on beer consumption. But what about the excise tax? Unlike ad valorem taxes, an excise tax is included as part of the label price.

Intuitively, do you expect consumers be more elastic towards changes in the ad valorem tax or excise tax?


Let's put this theory to the test. Each row in the table `excise_changes` displays an event in which there was a change in the excise tax, and its corresponding change in beer consumption. For example, from the table we can see that Arkansas has historically changed its beer excise tax three times. 

```python
# You do not need to understand how this code works.
excise_changes = pd.DataFrame({"st_name":[], "excise_tax_change":[], "c_beer_change":[]})
beer_df = beer.to_df()
for row_index in range(beer_df.shape[0]):
    if row_index == 0:
        continue
    if beer_df.iloc[row_index, 4] != beer_df.iloc[row_index-1, 4] and beer_df.iloc[row_index, 0] == beer_df.iloc[row_index-1, 0]:
        excise_changes = pd.concat([excise_changes, pd.DataFrame({"st_name": beer_df.iloc[row_index, 0], 
                                                "excise_tax_change": beer_df.iloc[row_index, 4] - beer_df.iloc[row_index-1, 4],
                                                "c_beer_change": beer_df.iloc[row_index, 7] - beer_df.iloc[row_index-1, 7]}, index=[row_index])], ignore_index=True)
excise_changes = Table.from_df(excise_changes[["st_name", "excise_tax_change", "c_beer_change"]])
excise_changes
```

```python
excise_changes.scatter(1,2,fit_line=True)
```

**Question 1.7:**
Using `np.polyfit`, determine the slope and intercept terms for this regression on the table `excise_changes`.

Recall that `np.polyfit` follows the syntax: `np.polyfit(<x array>, <y array>, <degree>)`.


```python
coeffs = ...
coeffs
```

```python
grader.check("q1_7")
```

## Question 2

*From Gruber, Ch.19, Q.4*

The demand for football tickets is $Q = 360 − 10P$ and the supply of football tickets is $Q = 20P$. 

**Question 2.1:**
Calculate the quantity and price at the equilibrium. Assign these values to `quantity_no_tax` and `price_no_tax`, respectively.


```python
P = sympy.Symbol("P")
supply = ...
demand = ...
price_no_tax = ...
quantity_no_tax = ...
print("The equilibrium price is", price_no_tax, "; the equilibrium quantity is", quantity_no_tax)
```

```python
grader.check("q2_1")
```

**Question 2.2:**
Assume that the government implements a per-ticket tax of $4 when consumers purchase a ticket. What is the new price received by producers and price paid by consumers?


```python
demand_with_tax = ...
price_received_with_tax = ...
price_paid_demand_with_tax = ...
print("The new price received by producers is", price_received_with_tax, 
      "; the new price paid by consumers is", price_paid_demand_with_tax)
```

```python
grader.check("q2_2")
```

**Question 2.3:**
Who bears a larger burden of the tax?

<ol style="list-style-type: lower-alpha;">
    <li>Producers</li>
    <li>Consumers</li>
</ol>

Assign `2_3` to letter corresponding to your choice.


```python
q2_3 = ...
```

```python
grader.check("q2_3")
```

**Question 2.4:**
Continue to assume that the government implements a per-ticket tax of \$4. How many less tickets are sold due to the tax?


```python
quantity_with_tax = ...
difference_in_quantity = ...
print("The difference in quantity transacted is", difference_in_quantity)
```

```python
grader.check("q2_4")
```

**Congratulations! You just finished Lab 3!**

Tests for q1_1 (points: [0])
```
>>> q1_1 in [1, 2, 3, 4]
True
```


Tests for q1_2 (points: None)
```
>>> len(beer) == 8
True
```
```
>>> "pc_beer" in beer.labels
True
```


Tests for q1_3 (points: None)
```
>>> len(np.unique(beer_state["st_name"])) == 1
True
```
```
>>> "CA" in np.unique(beer_state["st_name"])
True
```
```
>>> beer_state.num_rows == 34
True
```


Tests for q1_4 (points: None)
```
>>> len(beer_state_changes) == 2
True
```
```
>>> beer_state_changes.num_rows == 33
True
```
```
>>> for l in ["tax_diff", "consumption_diff"]:
...     assert l in beer_state_changes.labels
```


Tests for q1_5 (points: None)
```
>>> q1_5 in ["a", "b", "c", "d"]
True
```


Tests for q1_7 (points: None)
```
>>> len(coeffs) == 2
True
```
```
>>> -5 <= coeffs[0] <= 0 
True
```
```
>>> 0 <= coeffs[1] <= 1
True
```


Tests for q2_1 (points: None)
```
>>> 5 <= price_no_tax <= 15
True
```
```
>>> 200 <= quantity_no_tax <= 250
True
```


Tests for q2_2 (points: None)
```
>>> 10 <= price_received_with_tax <= 15
True
```
```
>>> 10 <= price_paid_demand_with_tax <= 15
True
```


Tests for q2_3 (points: None)
```
>>> q2_3 in ["a", "b"]
True
```


Tests for q2_4 (points: None)
```
>>> 25 <= difference_in_quantity <= 30
True
```
