```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

<table style="width: 100%;" id="nb-header>">
        <tr style="background-color: transparent;"><td>
            <img src="https://data-88e.github.io/assets/images/blue_text.png" width="250px" style="margin-left: 0;" />
        </td><td>
            <p style="text-align: right; font-size: 10pt;"><strong>Economic Models</strong>, Data 88E<br>
                Dr. Eric Van Dusen<br>
            Chris Pyles <br> Isabella Siu <br> Akhil Venkatesh</p></td></tr>
    </table>

# Lab 7: Game Theory

In this lab, we will recall some of the game theory concepts covered and review the prisoner's dilemma!

```python
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
import sympy
solve = lambda x,y: sympy.solve(x-y)[0] if len(sympy.solve(x-y))==1 else "Not Single Solution"
sympy.init_printing(print_builtin=False)
import pandas as pd
import datetime as dt
import seaborn as sns
%matplotlib inline
plt.style.use('seaborn-muted')
plt.rcParams['figure.figsize'] = [25,10]
from players import *
import warnings
warnings.simplefilter('ignore')
```

### Question 1

Consider the game defined by the payoff matrix below. In this game, the players cannot communicate when picking a strategy but know their opponent's history.

<table class="payoff-matrix" style="text-align: center; table-layout: fixed;">
    <tr style="background-color: #FFF;">
        <td></td>
        <td></td>
        <td style="text-align: center;" colspan="2">Player 2</td>
    </tr>
    <tr>
        <td></td>
        <td></td>
        <td width="85">Strategy A</td>
        <td width="85">Strategy B</td>
    </tr>
    <tr>
        <td rowspan="2" style="background-color: #FFF;">Player 1</td>
        <td>Strategy A</td>
        <td>State 1<br>(3, -2)</td>
        <td>State 2<br>(5, -5)</td>
    </tr>
    <tr>
        <td>Strategy B</td>
        <td>State 3<br>(-5, 5)</td>
        <td>State 4<br>(4, -1)</td>
    </tr>
</table>

**Question 1.1:** Assign `cooperative`, `symmetric`, `simultaneous`, and `perfect` below to `True` or `False` indicate whether that trait corresponds to this game. For example, if the game was assymetric, you would set `symmetric = False`.


```python
cooperative = ...
symmetric = ...
simultaneous = ...
perfect = ...
```

```python
grader.check("q1_1")
```

**Question 1.2:** Does this game have any Nash equilibria? List the numbers of the states in the payoff matrix above that correspond to Nash equilibria, if any, in the array `nash_equilibria` below. Leave the array empty if there are no Nash equilibria. For example, `make_array(1, 2)` means state 1 and 2 are Nash equilibria. 


```python
nash_equilibria = make_array(...)
nash_equilibria
```

```python
grader.check("q1_2")
```

### Question 2

In this section, we'll run through the application of the Cournot and Bertrand competition models. For this question, we'll be analyzing the rideshare duopoly of Uber and Lyft, who we will say share a marginal cost of $c$ = \\$3 per ride. Suppose the market demand is $P = -2.39 Q + 87.42$.

**Question 2.1:** Using SymPy, create a symbol for $Q$ and assign `market_demand` to the market demand curve for rideshares.


```python
Q = ...
market_demand = ...
market_demand
```

```python
grader.check("q2_1")
```

We'll start by assuming that Uber and Lyft form a Cournot duopoly and that Uber is trying to determine what output level to produce at under the assumption that Lyft will be producing at 15 units of output.

**Question 2.2:** Assign `profit_function` to Uber's profit function under the assumption that Lyft will produce 15 units of output. 

Hint: `q1` should represent the quantity produced by Uber (as the symbol $q_1$) and `q2` should represent the quantity produced by Lyft. `P` is the symbol for market price and `c` is the constant marginal cost mentioned before. The `profit_function` should be a function of $q_1$. 

```python
q_1 = ...
q_2 = ...
P = ...
c = ...

profit = P * q_1 - c * q_1
profit_function = profit.subs(..., market_demand.subs(..., ...))
profit_function
```

```python
grader.check("q2_2")
```

**Question 2.3:** Assign `q_1_star` to Uber's best response function based on `profit_function`. 

`q_1_star` should generate a numerical value equal to $q_1^*$. **Do not** calculate the answer by hand and then just write it in the cell below; instead, use Sympy to calculate the answer below.

Hint 1: Recall the first-order condition for firm 1: $\frac{\partial \pi_1}{\partial q_1} = 2mq_1 + mq_2 + b - c = 0$. (See [this textbook chapter](https://data-88e.github.io/textbook/content/07-game-theory/cournot.html) for a refresher)  
Hint 2: You can substitute value into `market_demand` to compute $mq_2 + b$. 


```python
q_1_star = ...
q_1_star
```

```python
grader.check("q2_3")
```

**Question 2.4:** Finally, given Uber's optimal output level $q_1^*$, calculate the price at which it will sell $p_1^*$.


```python
p_1_star = ...

print(f"Uber's output level: {q_1_star:.0f}")
print(f"Uber's price: ${p_1_star:.2f}")
```

```python
grader.check("q2_4")
```

Now let's examine Uber and Lyft as a Bertrand duopoly. The marginal cost and market demand will remain the same.

**Question 2.5:** Find the monopoly price and quantity `p_m` and `q_m` using SymPy. You can reuse the `market_demand` curve from earlier. Again, don't copy-paste the coefficients - use Sympy for your calculations.


```python
marginal_revenue = ...
q_m = ...
p_m = ...

print(f"Monopoly price: ${p_m:.2f}")
print(f"Monopoly quantity: {q_m:.0f}")
```

```python
grader.check("q2_5")
```

**Question 2.6:** Using the monopoly price and quantity, if Uber believes that Lyft will sell at \\$15 and we select $h=0.25$, what is the value of Uber's reaction function? Assign it to `uber_price`. No need to use Sympy for this question. 


```python
uber_price = ...
uber_price
```

```python
grader.check("q2_6")
```

We will be creating some custom Python classes to represent different playing strategies. We'll be using the `create_player_class` function (provided) to create these classes for us. Here are a few important things to know about the classes:

Consider a player class instance `player = Player()`. 
* `player.play()` is a method that represents a single move. It returns `True` if the player **defects** and `False` if they cooperate.
* `player.history` is an array of the previous moves of the player instance. `player.history.item(-1)`, for example, is the most recent move (the last element of the array).

Don't worry about the other methods that are defined for players; these are there for the code below to work but you don't need to concern yourself with them. Here's how the `create_player_class` function works:

1. Define a function, `f`, that takes two arguments: `self` and `opponent` (more about these later), and returns `True` if the player defects and `False` otherwise.
2. Call `create_player_class` with two arguments: a string for the class name, e.g. `"Player"`, and the play method `f`.

Let's illustrate this by creating `Defector`, a player that always defects. The `defector_play` function below always returns `True`, since the player always defects. We then create `Defector` using `create_player_class`.

### Question 3

```python
def defector_play(self, opponent):
    return True

Defector = create_player_class("Defector", defector_play)
```

Just like any class, we create a `Defector` instance by calling the constructor:

```python
Defector()
```

**Question 3.1:** Create a `Cooperator` class that always cooperates.

```
BEGIN QUESTION
name: q3_1
points:
    - 0
    - 0.25
    - 0.75
```

```python
def cooperator_play(self, opponent):
    return False # SOLUTION NO PROMPT
    """ # BEGIN PROMPT
    return ...
    """ # END PROMPT

Cooperator = create_player_class("Cooperator", cooperator_play) # SOLUTION
```

```python
## Test ##
cooperator_play(None, None) in [True, False]
```

```python
## Test ##
repr(Cooperator()) == "Cooperator"
```

```python
## Hidden Test ##
cooperator_play(None, None)
```

Now that we have players, how do we pit them against each other? Provided is a function called `payoff` that returns the values from our payoff matrix when provided two players by using their `play` methods. `payoff` returns a tuple of 2 elements corresponding to the payoff for the first and second player:

```python
>>> payoff(player_1, player_2)
(player 1 payoff, player 2 payoff)
```

```python
payoff(Defector(), Cooperator())
```

Our analysis of the prisoner's dilemma hinges on analyzing strategies over multiple rounds. For this reason, we need a way to pit players against each other and find out who has the fewest years accrued. The function `run_match` below will run a match of two players with `n` turns. If `winner` is `True`, it returns the winner of the match; if `winner` is `False`, it returns a list for each player containing the sequence of years accrued. Also note the first two lines of the function that reset the player history for each player. This is important because we might end up reusing instances of players, and we don't want histories from past matches to affect the results of this match.

```python
def run_match(p1, p2, n=5, winner=True):
    p1.reset_history()
    p2.reset_history()
    
    p1_years = make_array()
    p2_years = make_array()
    
    for i in range(n):
        p1_score, p2_score = payoff(p1, p2)
        p1_years = np.append(p1_years, p1_score)
        p2_years = np.append(p2_years, p2_score)
        
    if winner:
        p1_mean = np.mean(p1_years)
        p2_mean = np.mean(p2_years)

        if p1_mean < p2_mean:
            return p1
        else:
            return p2
    else:
        return p1_years, p2_years
```

For example, we can run a 10-turn match between a defector and cooperator with:

```python
run_match(Defector(), Cooperator(), n=10)
```

We see that the defector won the match, as expected. We can get the raw data (the number of years accrued by each player at each iteration) by setting `winner=False`:

```python
run_match(Defector(), Cooperator(), n=10, winner=False)
```

Now that we know how to run matches, let's think of some other cool strategies. Remember the `player.history` variable discussed earlier? Let's have a concrete example on how to use that. The `Alternator` player defined below alternates between cooperating and defecting, beginning with cooperating.

Note the `if` statement. We need to check that we have a history to index into before we try to grab the last element; if we didn't have that statement, the code would fail on the first time  `Alternator` played, because we would try to get an element from an array with 0 elements.

```python
def alternator_play(self, opponent):
    if len(self.history) > 0:
        return not self.history.item(-1)
    else:
        return False
    
Alternator = create_player_class("Alternator", alternator_play)

run_match(Alternator(), Defector())
```

**Question 3.2:** Use the player history to define `TitForTat`, a player that does the last thing it's opponent did and starts off by cooperating. For example, if we pitted `TitForTat` against a `Defector` in a 5-round match, the `Defector`'s moves would be `[D, D, D, D, D]`, and `TitForTat`'s would be `[C, D, D, D, D]`.

```
BEGIN QUESTION
name: q3_2
points: 3
```

```python
def tit_for_tat_play(self, opponent):
    # BEGIN SOLUTION NO PROMPT
    if len(opponent.history) > 0:
        return opponent.history.item(-1)
    else:
        return False
    # END SOLUTION
    """ # BEGIN PROMPT
    if len(opponent.history) > 0:
        return ...
    else:
        return ...
    """ # END PROMPT

TitForTat = create_player_class("TitForTat", tit_for_tat_play) # SOLUTION
```

```python
## Test ##
repr(TitForTat()) == "TitForTat"
```

```python
## Test ##
d = Defector()
t = TitForTat()
tit_for_tat_play(t, d) in [True, False]
```

```python
## Test ##
d = Defector()
t = TitForTat()
_ = payoff(d, t)
tit_for_tat_play(t, d) in [True, False]
```

```python
## Hidden Test ##
d = Defector()
t = TitForTat()
tit_for_tat_play(t, d)
```

```python
## Hidden Test ##
d = Defector()
t = TitForTat()
_ = payoff(d, t)
tit_for_tat_play(t, d)
```

```python
## Hidden Test ##
c = Cooperator()
t = TitForTat()
_ = payoff(c, t)
tit_for_tat_play(t, c)
```

**Question 3.3:** Suppose we have a 10-turn match with a backstabber and a desperate. Write down the sequence of defections (D) and cooperations (C) for each player. Who wins the match?

```
BEGIN QUESTION
name: q3_3
points:
    - 0
    - .75
    - .75
    - .5
```

```python
# BEGIN SOLUTION NO PROMPT
backstabber_moves = make_array("C", "C", "C", "C", "D", "D", "D", "D", "D", "D")
desperate_moves = make_array("D", "D", "D", "D", "D", "C", "C", "C", "C", "C")
backstabber_desperate_winner = "backstabber"
# END SOLUTION
""" # BEGIN PROMPT
backstabber_moves = make_array(...)
desperate_moves = make_array(...)
backstabber_desperate_winner = "..."
"""; # END PROMPT
```

```python
## Test ##
assert all(i in ["C", "D"] for i in backstabber_moves)
assert all(i in ["C", "D"] for i in desperate_moves)
assert len(backstabber_moves) == 10
assert len(desperate_moves) == 10
assert backstabber_desperate_winner.lower() in ["backstabber", "desperate"]
```

```python
## Hidden Test ##
list(backstabber_moves) == ["C", "C", "C", "C", "D", "D", "D", "D", "D", "D"]
```

```python
## Hidden Test ##
list(desperate_moves) == ["D", "D", "D", "D", "D", "C", "C", "C", "C", "C"]
```

```python
## Hidden Test ##
backstabber_desperate_winner.lower() == "backstabber"
```

### Question 4

### Tournaments

Now that we have re-familiarized ourselves with our prisoner's dilemma code from the lab, let's move on to recreating Axelrod's tournament. Recall that Axelrod originally wanted to compare strategies for playing the prisoner's dilemma, which he did by pitting every strategy against every other strategy in a round-robin tournament. In this project, we will do the same thing, albeit on a much smaller scale, with the strategies defined in the table above.

We define the function `run_tournament` below that takes in an array of players and returns a table with four columns: `p1`, `p1_mean`, `p2`, and `p2_mean`. As in Axelrod's original tournament, we use matches with 200 rounds.

```python
def run_tournament(players):
    p1 = make_array() #This stores an array of all player 1's.
    p2 = make_array() #This stores an array of all player 2's.
    p1_means = make_array() #This stores an array of each player 1's average jail time for a particular game.
    p2_means = make_array() #This stores an array of each player 2's average jail time for a particular game.

    for i in range(len(players)):
        for j in range(i+1, len(players)):
            p1_years, p2_years = run_match(players.item(i), players.item(j), n=200, winner=False)
            p1 = np.append(p1, players.item(i))
            p2 = np.append(p2, players.item(j))
            p1_means = np.append(p1_means, np.mean(p1_years))
            p2_means = np.append(p2_means, np.mean(p2_years))
                
    results = Table().with_columns(
        "p1", p1,
        "p1_mean", p1_means,
        "p2", p2,
        "p2_mean", p2_means
    )

    return results

run_tournament(make_array(Defector(), Cooperator()))
```

**Question 4.1:** Create a tournament with all players in the table above Question 1.6. Assign the results table to `tournament_results`.

```
BEGIN QUESTION
name: q4_1
points:
    - 0
    - 0
    - 0.5
    - 0.5
    - 1
```

```python
# BEGIN SOLUTION NO PROMPT
all_strategies = make_array(
    Alternator(), Backstabber(), Bully(), Desperate(), FoolMeOnce(), Forgiver(), 
    ForgivingTitForTat(), Grudger(), OnceBitten(), TitForTat())
tournament_results = run_tournament(all_strategies)
# END SOLUTION
""" # BEGIN PROMPT
all_strategies = make_array(...)
tournament_results = run_tournament(...)
"""; # END PROMPT
tournament_results
```

```python
## Test ##
tournament_results.num_rows == 45
```

```python
## Test ##
tournament_results.num_columns == 4
```

```python
## Test ##
players = make_array(Alternator(), Backstabber(), Bully(), Desperate(), FoolMeOnce(), Forgiver(), 
                     OnceBitten(), TitForTat(), Grudger(), ForgivingTitForTat())
for player in players:
    assert player in tournament_results.column("p1") or player in tournament_results.column("p2")
```

```python
## Test ##
t = tournament_results.where("p1", Alternator()).where("p2", Bully())
assert np.isclose(t["p1_mean"][0], 3.015)
assert np.isclose(t["p2_mean"][0], 2.99)
```

```python
## Hidden Test ##
t = tournament_results.where("p1", Desperate()).where("p2", Grudger())
assert np.isclose(t["p1_mean"][0], 4.97)
assert np.isclose(t["p2_mean"][0], 0.045)
```

To visualize our results, we will make a [heat map](https://en.wikipedia.org/wiki/Heat_map) that displays our results in an easy-to-read manner. The code in the cell below uses seaborn and pandas, two libraries you will learn about in Data 100, so don't worry about understanding it.

The plot shows Player 1 on the horizontal axis, Player 2 on the vertical axis, and the values in the cells are **Player 1's mean score** (i.e. the mean score of the player on the horizontal axis). So, if we look at the intersection of `OnceBitten` on the horizontal and `Forgiver` on the vertical, we see `OnceBitten`'s mean of 2.000 and if we do the reverse (`Forgiver` on horizontal, `OnceBitten` on vertical), we see `Forigiver`'s mean of 2.000.

<!-- BEGIN QUESTION -->

```python
df1 = tournament_results.to_df()
df2 = tournament_results.to_df()
df2 = df2.rename({"p1": "p2", "p2": "p1", "p1_mean": "p2_mean", "p2_mean": "p1_mean"}, axis=1)
df = df1.append(df2)
df["p1"], df["p2"] = df["p1"].astype(str), df["p2"].astype(str)
df = df.pivot("p1", "p2", "p1_mean")

plt.figure(figsize=[15,8])
sns.heatmap(df.T, annot=True, fmt=".3f")
plt.xlabel("Player 1")
plt.ylabel("Player 2")
plt.title("Player 1's Mean Scores by Opponent")
plt.xticks(rotation=90);
```

**Non-Graded Reflection Question** Interpret the plot above. What results stand out to you? Which strategy would you choose if you were playing in a tournament? Justify your answer.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> for var in [cooperative, symmetric, simultaneous, perfect]:
...     assert var in [True, False]
```


Tests for q1_2 (points: [0, 0])
```
>>> assert 0 <= len(nash_equilibria) <= 4
```
```
>>> for val in nash_equilibria:
...     assert val % 1 == 0
```


Tests for q2_1 (points: None)
```
>>> assert type(Q) == sympy.Symbol
>>> assert Q in market_demand.as_coefficients_dict()
>>> assert 1 in market_demand.as_coefficients_dict()
```


Tests for q2_2 (points: 2)
```
>>> assert type(q_1) == sympy.Symbol, "q_1 should be a Sympy symbol"
>>> assert type(q_2) == int, "q_2 should be a number"
>>> assert type(P) == sympy.Symbol, "P should be a Sympy symbol"
>>> assert type(c) == int, "c should be a number"
>>> assert q_1 in profit_function.as_coefficients_dict(), "q_1 should be a variable in profit_function"
>>> assert P not in profit_function.as_coefficients_dict(), "P cannot be a variable in profit_function"
>>> assert len(profit_function.as_coefficients_dict()) == 2
```


Tests for q2_3 (points: [0])
```
>>> assert type(q_1_star) == sympy.core.numbers.Float, "You need to use Sympy"
>>> assert 0 < float(q_1_star) < 20
```


Tests for q2_4 (points: None)
```
>>> assert type(p_1_star) == sympy.core.numbers.Float, "You need to use Sympy"
>>> assert 20 < float(p_1_star) < 40
```


Tests for q2_5 (points: [0, 0])
```
>>> assert 35 < float(p_m) < 50
```
```
>>> assert 10 < float(q_m) < 20
```


Tests for q2_6 (points: None)
```
>>> assert 3 <= uber_price <= 45.21
```
